rule_files:
  - ../../../src/prometheus_alert_rules/ipmi_sel.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: ipmi_sel_command_success{instance="ubuntu-0"}
        values: '0x15'

      - series: ipmi_sel_command_success{instance="ubuntu-nominal"}
        values: '1x15'
      - series: ipmi_sel_state_nominal{instance="ubuntu-nominal"}
        values: '1x5 2x5 3x5'

      - series: ipmi_sel_command_success{instance="ubuntu-warning"}
        values: '1x15'
      - series: ipmi_sel_state_warning{instance="ubuntu-warning"}
        values: '1x5 2x5 3x5'

      - series: ipmi_sel_command_success{instance="ubuntu-critical"}
        values: '1x15'
      - series: ipmi_sel_state_critical{instance="ubuntu-critical"}
        values: '1x5 2x5 3x5'

      - series: node_systemd_unit_state{name="ipmiseld.service", instance="ubuntu-3", state="failed"}
        values: '1x15'

      - series: node_systemd_unit_state{name="ipmiseld.service", instance="ubuntu-4", state="inactive"}
        values: '1x15'

    alert_rule_test:
      - eval_time: 10m
        alertname: IPMISELCommandFailed
        exp_alerts:
          - exp_labels:
              severity: critical
              instance: ubuntu-0
            exp_annotations:
              summary: Failed to run ipmi-sel. (instance ubuntu-0)
              description: |
                Failed to get system event logs using ipmi-sel.
                  VALUE = 0
                  LABELS = map[__name__:ipmi_sel_command_success instance:ubuntu-0]

      - eval_time: 7m
        alertname: IPMISELStateWarning
        exp_alerts:
          - exp_labels:
              severity: warning
              instance: ubuntu-warning
              event_id: 2
            exp_annotations:
              summary: IPMI system event log in warning state. (instance ubuntu-warning)
              description: |
                IPMI SEL entry in warning state.
                  LABELS = map[__name__:ipmi_sel_state_warning instance:ubuntu-warning]
                  EVENT_ID = 2

      - eval_time: 13m
        alertname: IPMISELStateWarning
        exp_alerts:
          - exp_labels:
              severity: warning
              instance: ubuntu-warning
              event_id: 3
            exp_annotations:
              summary: IPMI system event log in warning state. (instance ubuntu-warning)
              description: |
                IPMI SEL entry in warning state.
                  LABELS = map[__name__:ipmi_sel_state_warning instance:ubuntu-warning]
                  EVENT_ID = 3

      - eval_time: 7m
        alertname: IPMISELStateCritical
        exp_alerts:
          - exp_labels:
              severity: critical
              instance: ubuntu-critical
              event_id: 2
            exp_annotations:
              summary: IPMI system event log in critical state. (instance ubuntu-critical)
              description: |
                IPMI SEL entry in critical state.
                  LABELS = map[__name__:ipmi_sel_state_critical instance:ubuntu-critical]
                  EVENT_ID = 2

      - eval_time: 13m
        alertname: IPMISELStateCritical
        exp_alerts:
          - exp_labels:
              severity: critical
              instance: ubuntu-critical
              event_id: 3
            exp_annotations:
              summary: IPMI system event log in critical state. (instance ubuntu-critical)
              description: |
                IPMI SEL entry in critical state.
                  LABELS = map[__name__:ipmi_sel_state_critical instance:ubuntu-critical]
                  EVENT_ID = 3

      - eval_time: 10m
        alertname: IPMISELDStateWarning
        exp_alerts:
          - exp_labels:
              severity: warning
              name: ipmiseld.service
              instance: ubuntu-3
              state: failed
            exp_annotations:
              summary: IPMISELD service is not active. (instance ubuntu-3)
              description: |
                The ipmiseld service is not active, indicating a potential problem.
                  VALUE = 1
                  LABELS = map[__name__:node_systemd_unit_state instance:ubuntu-3 name:ipmiseld.service state:failed]

          - exp_labels:
              severity: warning
              name: ipmiseld.service
              instance: ubuntu-4
              state: inactive
            exp_annotations:
              summary: IPMISELD service is not active. (instance ubuntu-4)
              description: |
                The ipmiseld service is not active, indicating a potential problem.
                  VALUE = 1
                  LABELS = map[__name__:node_systemd_unit_state instance:ubuntu-4 name:ipmiseld.service state:inactive]
